---
title       : chapter1.md
description : Insert the chapter description here
--- type:NormalExercise lang:r xp:100 skills:1 key:afd887fe02
## <<<New Exercise>>>


*** =instructions
Aqui deve existir um texto de instrucoes

*** =hint
o que esta nesta parte ficará escondiida

*** =pre_exercise_code
```{r}
x<-2
```

*** =sample_code
```{r}
x<-2
print(x)
```

*** =solution
```{r}
x<-2
y<-3
z<-x+y
```

*** =sct
```{r}
x<-2
y<-3
z<-x+y
print(x)
```
