---
title       : Insert the chapter title here
description : Insert the chapter description here
--- type:NormalExercise lang:r xp:100 skills:1 key:916b1c3e9a
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```


1) Qual a dimensão do dataset Iris?

```{r}

```


2) Qual o nome dos atributos presentes no dataset?

```{r}

```

3) Qual o tipo de cada atributo?

```{r}

```

4) Apresente um resumo estatístico do dataset Iris e faça algumas análises em termos de distribuição e intervalo de valores.

```{r}

```
