---
title       : Teste de gamfication
description : Este capitulo foca a gamification

--- type:NormalExercise lang:r xp:100 skills:1,4 key:d3148b4c04

--- type:MultipleChoiceExercise lang:r xp:50 skills:3 key:e74c30a746
## 2222 Multiple Choice Exercise Title

2222 This is the assignment.
How much is `2 + 2`?

*** =instructions
- 2
- 3
- 4

*** =hint
Here's a hint: think about your time in primary school.

*** =pre_exercise_code
```{r}
# no pec
```

*** =sct
```{r}
msg1 = "Try again! Think about your time in school."
msg2 = "Summing two even numbers always leads to another even number. Try again."
msg3 = "Well done. Proceed to the next exercise"
test_mc(correct = 3, feedback_msgs = c(msg1,msg2,msg3))
```

... _rest of exercise omitted_
